```python
#this is to plot inside the notebook
%matplotlib inline

# Let's start by importing the libraries you will be using in this assignment
import numpy as np
import sklearn
import matplotlib.pyplot as plt
import matplotlib

# A seed for the random number generator so all results are reproducible
np.random.seed(15)

# Makes our plots look nicer
matplotlib.style.use("seaborn-notebook")

# Let's also hide unnecessary warnings
import warnings
warnings.filterwarnings('ignore')
```

# Assignment 1 Set 2: Learning Categories for Mammals and Birds

## A Perceptron for Mammals and Birds
This part of the assignment is designed to look at neural networks in the context of a more complicated classification problem, and to look at a data set related to human concepts and categories.


### Classifying animals

We are going to figure out whether a perceptron can learn to distinguish animals into two classes: mammals and birds. In the data folder, there should be 9 different files. 

Load the data in the notebook:

```python
# import the data to train the network on
birds = np.genfromtxt(open('./data/birds.txt', "rb"), delimiter=",").T
mammals = np.genfromtxt(open('./data/mammals.txt', "rb"), delimiter=",").T

birds_tok = np.genfromtxt(open('./data/birds_tok.txt', "rb"),dtype='str', delimiter=",")
mammals_tok = np.genfromtxt(open('./data/mammals_tok.txt', "rb"),dtype='str', delimiter=",")

feats = np.genfromtxt(open('./data/feats.txt', "rb"),dtype='str', delimiter=",")
```

To train the perceptron, we will be using data from [McRae et al. (2005)(1)](#references) . You may want to have a quick read through the paper before doing the assignment.

To inspect the data you can use: 
    
`print(filename)` 

The **birds** and **mammals** matrices should have 20 rows and 16 columns. Each column in the matrix corresponds to a visual feature (e.g., has legs, flies) and each row to an animal (e.g., stork, pigeon). The features are 16 in total, and the animals are 20 per matrix (i.e., 20 birds and 20 mammals).

Inspect the different files in the notebook, and make sure that they are in the correct dimensions (using .shape).

( 2 points)

<!-- grade_id: cell-406b5ea33d1e6b7d, points: 1, grade: true, solution: true -->
```python
# inspect files here

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8c096b338c9f7ab1, points: 0, grade: false, solution: true -->
```python
# compute the shapes of the birds and mammals matrices here
# birds_shape =  ...      
# mammals_shape = ...  
# animal_tokens = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1634927a39df7c53, points: 1, grade: true, solution: false -->
```python
"""Check that *mammals* and *birds* matrices are the  correct shape"""
assert(birds_shape == mammals_shape)
assert(mammals_shape == (20,16))
```

The features are defined in the vector `feats`, the birds are defined in vector `birds_tok`, and the mammals in vector `mammals_tok`. 

The visual features were elicited from humans who were presented with a word (e.g., *tiger*) and asked to list its properties. The cells in the mammals and birds matrices correspond to the number of humans who listed a specific feature for each word. The counts are normalized so that each matrix row sums to one.

The goal is to train a perceptron that distinguishes between mammals and birds, by deciding whether a given input is a bird. In other words, the perceptron should output 1 in the case of a bird and 0 in the case of a mammal. You will need a matrix, input, containing all training examples, and a target vector, targets birds, whose entry i is 1 if example i is a bird, and 0 otherwise (use `np.zeros(n)` and `np.ones(n)`) for this.

For this, use `np.concatenate((matrix1,matrix2))` and use `np.zeros(n)` and `np.ones(n))` for the target vector. 

(1 point)

<!-- grade_id: cell-53fb5c06c44e1bfe, points: 0, grade: false, solution: true -->
```python
# create the training data matrix X_train and target vector y_train
# X_train =  ...   # concatenate birds and mammals
# y_train =  ...   # 1 if bird, 0 if mammal 
# tokens_train = ... # concatenate animal names

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7fb2103dc809502c, points: 1, grade: true, solution: false -->
```python
"""Check that the training and target matrices are in the correct shape"""
assert(np.shape(X_train) == (40,16))

```

Now create the perceptron, similar to how we created the MLP and linear regression in assignment 1 part 1. Create a new perceptron and fit it to the data like we did in task 1.

(2 points)


<!-- grade_id: cell-fa53d1a3d53a3268, points: 2, grade: true, solution: true -->
```python
#import the perceptron model from scikit-learn
from sklearn.linear_model import Perceptron

# Create Perceptron 
# net = ...

# Fit network to the data X_train and y_train

# Make predictions on training data 
# predictions = ...

# YOUR CODE HERE
raise NotImplementedError()
```

Report the categorization error of the perceptron on the training data using the mean_squared_error function imported in the cell below. 

(2 points)

<!-- grade_id: cell-eaab20dc1f5e9c01, points: 0, grade: false, solution: true -->
```python
# import mean squared error from scikit-learn
from sklearn.metrics import mean_squared_error
# calculate MSE here
# MSE = mean_squared_error(... , ...)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-047a926eb4383e53, points: 2, grade: true, solution: false -->
```python
''' check for MSE value'''
assert MSE is not None

```

Did the model classify all examples correctly? You can see which examples were classified incorrectly by looking at the differences between the `.predict(data)` array and the target array.

(2 points)

<!-- grade_id: cell-7fcb3a7378ec99f0, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Briefly discuss the classification results. 

(3 points)

<!-- grade_id: cell-7c6dc16c50703421, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

# Comparing different models.

Now we'll be using a slightly more complicated neural network, resembling the one we saw in Tutorial 3. It has more units in its hidden layer, but otherwise it's very similar. It might be helpful to know that unlike our in-class example, the weights don't always start at zero or another fixed value. Like with the perceptron, fit the new network to the training data and assess whether it's learned the correct classification by making it predict the categories of the animals it's been trained on.


(3 points)

<!-- grade_id: cell-c34d2ec827afbe37, points: 3, grade: true, solution: true -->
```python
#Let's create the network - we use the logistic sigmoid function (logistic) as activation function
# and stochastic gradient descent ('sgd') to train the network
from sklearn.neural_network import MLPClassifier

net2 = MLPClassifier(activation='logistic', solver='sgd', tol=0,learning_rate_init=.08, max_iter=200)

# fit network here

# predict categories of training data here

# YOUR CODE HERE
raise NotImplementedError()
```

Calculate the MSE of the new network and examine its performance.

(2 points)

<!-- grade_id: cell-e0c071562f56ee6c, points: 0, grade: false, solution: true -->
```python
# MSE_result = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-65e4ebb628d5e4e8, points: 2, grade: true, solution: false -->
```python
'''test that MSE result is not NoneType'''
assert MSE_result is not None

```

 Did the model learn the correct classification with the set parameters? Discuss. 
 
 (2 points)

<!-- grade_id: cell-c13a9d6240de97b9, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Now, recreate and re-fit the model ten times to the training data (`using net2.fit` and `net2.predict`), and compare the performances. To do this, store the MSE of each run in an array.

(2 points)

<!-- grade_id: cell-2e01cb14f8cde174, points: 0, grade: false, solution: true -->
```python
# MSE_runs = []  # create an array to store each result

# for i in range(10):  # iterate 10 times
    # net2 = ...              recreate model
    #                         fit model to training data (X_train, y_train)
    # net_2_predictions = ... predictions
    
    # MSE = ...               compute MSE
    # MSE_runs.append(MSE)  # this stores the result of one run to the MSE_runs array 
    
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-5de25dbf821a2368, points: 2, grade: true, solution: false -->
```python
''' check length of MSE array'''
assert len(MSE_multiple_runs)==10
```

Make a bar plot of all 10 MSEs here using `plt.bar`. Note that `plt.bar` needs two arrays: one for the indices, and one for the values. For example, to plot `[1,2,3,5,8]`:

    plt.bar(np.arange(5), [1,2,3,5,8]) 

Remember to add title and x- and y-labels.

(3 points)

<!-- grade_id: cell-de771037eb9d217b, points: 3, grade: true, solution: true -->
```python
# plot below

# YOUR CODE HERE
raise NotImplementedError()
```

Briefly discuss your results. Would you observe the same results with the perceptron?

(3 points)

<!-- grade_id: cell-e1575069040c750a, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Generalizing  to new data

Let us now see how well the perceptron can generalize to unseen animals. Generalization happens when a model re-uses what it has learned in the training data to examples it did not see during training. 

You will have to prepare the test data in the same way as the training data: you will again need a matrix with all test examples (the input to the perceptron), and a target vector. Make sure that your use of zeros and ones matches mammals and birds in the same way that it did during the training phase!

( 2 points)

```python
#import test data here
mammals_test = np.genfromtxt(open('./data/mammals_test.txt', "rb"), delimiter=",").T
mammals_tok_test = np.genfromtxt(open('./data/mammals_tok_test.txt', "rb"),dtype='str', delimiter=",")

birds_test = np.genfromtxt(open('./data/birds_test.txt', "rb"), delimiter=",").T
birds_tok_test = np.genfromtxt(open('./data/birds_tok_test.txt', "rb"),dtype='str', delimiter=",")
```

<!-- grade_id: cell-98d01248ac1a7b8c, points: 0, grade: false, solution: true -->
```python
# create the training data matrix X_train, target vector y_train and a vector with the names of the animals
# X_test =  ...  
# y_test =  ...   
# test_tokens = ...


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-74f6fd7cf316ac4a, points: 2, grade: true, solution: false -->
```python
'''check shapes for test data '''
assert(X_test.shape == (30,16))

```

Now test the perceptron without re-training it on the data (i.e. using `net1.predict(X_test)`). This means that the network is now predicting the category of mammals and birds that it has not seen.

(2 points)

<!-- grade_id: cell-e4d608f9975821b5, points: 0, grade: false, solution: true -->
```python
# net1_test_predictions = ...
# test_MSE = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5ed839c5cd39532d, points: 2, grade: true, solution: false -->
```python
'''check test set predictions'''
assert(len(net1_test_predictions) == len(X_test))

```

Discuss the results briefly (one or two sentences).
 
(2 points)

<!-- grade_id: cell-378e2ca9e4639e2e, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Inspecting model predictions.
Which animals did the neural network not classify correctly? Complete the function below so that it prints the misclassified animals and the weight attributed to each feature.

( 2 points)

<!-- grade_id: cell-24f8a69c852d59c1, points: 2, grade: true, solution: true -->
```python
for index, true_label in enumerate(y_test):
    if net1_test_predictions[index]!=true_label: # if prediction does not match the true category
        # here, print the name of the animal corresponding to the index
        
        # below, print the feature weights of the animal corresponding to the index
        
        # YOUR CODE HERE
        raise NotImplementedError()
```

Can you find an explanation for that? Looking at the weights that the network learned in conjunction with the McRae features (corresponding to the `feats` file) the weights operate on might help you answer the question. 

First print out the weights of the perceptron by using the `net1.coef_` function.

(1 point)

<!-- grade_id: cell-573032e5fce7bdb9, points: 1, grade: true, solution: true -->
```python
# code here

# YOUR CODE HERE
raise NotImplementedError()
```

Discuss the animals that the model failed to classify correctly using the learned weights of the model and the features of the misclassified animals.

(5 points)

<!-- grade_id: cell-3792a6eeb2e17334, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Comparing models on the test set.
Now, calculate the MLP (`net2`) error on the test set. How does it compare to the single layer perceptron performance?

Frist print out the MSEs for `net2` below:

(3 points)

<!-- grade_id: cell-202e874cf28a4087, points: 3, grade: true, solution: true -->
```python
# net2_test_predictions = ...

# Calculate MSE 

net2_test_predictions = net2.predict(X_test)
print (mean_squared_error(net2_test_predictions, y_test))

for index, true_label in enumerate(y_test):
    if net2_test_predictions[index]!=true_label: # if prediction does not match the true category
        # here, print the name of the animal corresponding to the index
        
        # below, print the feature weights of the animal corresponding to the index
        
        # YOUR CODE HERE
        raise NotImplementedError()


```

Now discuss the MSEs, the number of erroneous classifications, and the differences/similarities between animals the network was not able to classify.

<!-- grade_id: cell-c33635af14208403, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Reliability 

Do ten iterations, where in each iteration you create the network (net2) again, train it and examine its performance on the training and test data. Remember to NOT fit the network on the test data. With the test set, the model is meant to predict the category of animals it has never seen before (i.e. animals it has not been trained on).

On the same graph, plot the performance of each iteration on the test set (using `plt.bar`). Don’t forget to add labels to both axes.

You can plot the test error on top of the train error in your bar plots by passing the parameter `bottom` in the `plt.bar()` function. Add separate labels for training and test error. 

(3 points)

<!-- grade_id: cell-dbeae86f309398d9, points: 3, grade: true, solution: true -->
```python
# code here

# MSE_train = []
# MSE_test = []
# for i in range(10):
#   ...
#   ...
  
    
# YOUR CODE HERE
raise NotImplementedError()
```

Discuss your observations briefly.

(2 points)

<!-- grade_id: cell-331b9ee678ec5600, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Playing with parameters 
Experiment with the following parameters:
- Learning rate: change the learning rate settings  ($\texttt{learning_rate_init: e.g. `1, 0.1,0.01,0.001`) 
- Number of hidden layers and units vary the number of units in the hidden layer, and try to add a second hidden layer. (e.g. `(25,), (80,),(100,), (100,80), (100,100)`)
- Iterations: change the max number of iterations. E.g. `20,100,500,1000, 5000`

Follow the same procedure as before, running each simulation 10 times (for each new parameter).

**Notes**:


- To compare different parametrizations, you should plot the mean and standard error for the 10 runs of each model with a given parametrization. You can use np.mean np.std to compute these. You plot this through the `yerr` parameter in the `plt.bar()` function. You should pass two arrays to the plt.bar function: the means and the std  of different parametrizations.

- Make three plots. In each, vary either the learning rate, the number of hidden layers and the iterations.

(10 points) 

<!-- grade_id: cell-6e9fd56756fe0075, points: 4, grade: true, solution: true -->
```python
#code here for plot 1 (learning rate)  -- 4 pts

# MSE_mean = []
# MSE_se = []
# for lr in [1,0.1,0.01,0.001,0.0001]:
#    model_results = []  # an empty array is re-created for every learning rate value
#    for i in range(10):
        # create network with given learning rate
       
        # fit network to training data
        
        # predict on test data 
        
        # compute MSE of predictions on test data
        
        # append to model_results
    
    # at the end of the for loop 
    # append the mean to MSE_mean array
    # append the std to MSE_se array

# Make plot here
    
# to re-name the tick labels for each bar, use:    
# plt.xticks(np.arange(5), ['1', '0.1','.01', '.001', '.0001']);


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1bab9c9c613cfa07, points: 3, grade: true, solution: true -->
```python
#code here for plot 2 (hidden layers)  -- 3 pts
# MSE_mean = []
# MSE_se = []

# for hl in [(25,), (80,),(100,), (100,80), (100,100)]:
#     model_results = []
#     for i in range(10):
        
#         # create network with given hidden layers
#         net2 = MLPClassifier(activation='logistic', solver='sgd',hidden_layer_sizes= hl,
#                              learning_rate_init=0.1, max_iter=250)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cbdba4471424bcd1, points: 3, grade: true, solution: true -->
```python
#code here for plot 3 (max-iter)  -- 3 pts

# MSE_mean = []
# MSE_se = []
# for it in [20,100,500,1000, 5000]: 
#     model_results = []
#     for i in range(10):
        
#         # create network with given hidden layers
#         net2 = MLPClassifier(activation='logistic', solver='sgd', tol=0,
#                              learning_rate_init=0.08, max_iter=it)


# YOUR CODE HERE
raise NotImplementedError()
```

What is the best result you get on the training and test set, respectively? What were the corresponding parameter settings? Discuss your results and the effect of each parameter on the performance of the model.

You might want to have a look at Chapter 6 in "An Introduction to Neural Networks" (Gurney, 1997) for information.

(8 points)

<!-- grade_id: cell-5d03c43438d63b70, points: 8, grade: true, solution: true -->
YOUR ANSWER HERE

# <a id='references'>References</a> 
[1] [K. McRae, G.S. Cree, M.S. Seidenberg, C. McNorgan. 2005. Semantic Feature Production Norms for a Large Set of Living and Nonliving Things. Behavior Research Methods, 37(4):547-59.](https://link.springer.com/content/pdf/10.3758/BF03192726.pdf)
